---
title: 'Discrete Distributions using R'
description: 'This lab is about working with discrete distributions using R'
attachments: null
---

## <<<New Exercise>>>

```yaml
type: MultipleChoiceExercise
key: f68b666117
lang: r
xp: 50
skills: 1
```



`@possible_answers`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sct`
```{r}

```
