---
title       : Using ggplot2 for ACTS 121
description : This lab is about learning to use ggplot2 for graphics in ACTS 121
--- type:NormalExercise lang:r xp:100 skills:1 key:66550becfc
## European Call Option (part 1)

A European call option on stock index $S$ has strike price $K = 40$ and will expire at time $T$.  In this series of exercises you will learn how
to use ggplot2 to create an attractive graph of the option's payoff at expiration.  The graphs you learn to create can be included in your Rmarkdown files or power point presentation slides.  In part 1 of this series you will set up the payoff function and create a data frame that will be used for plotting.



*** =instructions

* Define the option call option payoff function as `call_payoff()`
* `call_payoff()` should have two arguements:
* - $S$, the stock price at expiration
* - $K$, the strike price
* - Let $K$ have the default value $K = 40$

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
