---
title       : Discrete Distributions using R
description : This lab is about working with discrete distributions using R
attachments :


--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:f68b666117
## <<<New Exercise>>>


*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}

```
